--- 
title       : "Regressions 2: Introduction to Matching Methods"
description : "This chapter will introduce you to using matching methods to find causal effects"
 
 
--- type:VideoExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:149c81ed6d
## Matching Methods
*** =video_link
//player.vimeo.com/video/217555077

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:ffc8655145
## Reasons to Use Matching Methods. 
Megan the exercise freak (health coach) believes that eating out at restaurants can be good for people because it motivates them to exercise more the following day. She asks thousands of people whether they ate out at a restaurant the previous Saturday, and whether they spent any time exercising the following Sunday. She finds a positive correlation between eating at a restaurant and exercising the following day. Why might why might she want to use matching methods to assess whether there is a causal effect?

*** =instructions
- Because you didn't ask whether people why they chose to exercise on Sunday.
- Because most fitness clubs meet on Sundays.
- Because Sunday is most people's favorite day to exercise.
- Because people who eat out might have a higher propensity to exercise as well.

*** =sct
```{r}
msg1 = "Try again"
msg2 = "Try again"
msg3 = "Well done"
msg4 = "Try again"
test_mc(correct = 4, feedback_msgs = c(msg1,msg2,msg3,msg4))
```




--- type:VideoExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:d757f83840
## The Lifetime Earnings of Veterans and Nonveterans
*** =video_link
//player.vimeo.com/video/217555250


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:59172334ed
## Why we Need Matching Methods
Why couldn't Angrist just compare the lifetime earnings of veterans and nonveterans with standard regression methods?
*** =instructions
- Because the people who applied to the military and got rejected are probably poor people in the first place.
- Because the military only rejects disabled people.
- Because the veterans who selected into the military may have had very different earnings trajectories than individuals who had not selected into the military.
- Because people who got rejected are likely to be disproportionately female.
*** =sct
```{r}
msg1 = "Correct!"
msg2 = "Try again"
msg3 = "Well done"
msg4 = "Try again"
test_mc(correct = 1, feedback_msgs = c(msg1,msg2,msg3,msg4))
```



--- type:VideoExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:1e1cf83d18
## The Unconfoundedness Assumption in Angrist's Study on the Effect of Veteran Status on Lifetime Earnings
*** =video_link
//player.vimeo.com/video/217555486



--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:0cb4870c69
## Replication and Validity
Suppose you replicate the Angrist analysis using more recent data from 1992-2015, and find approximately zero causal effects for both whites and nonwhites, for all stages of life. On its own, would this finding invalidate Angrist's analysis?
*** =instructions
- Yes
- No
*** =sct
```{r}
msg1 = "Try again"
msg2 = "Correct! This is a relatively difficult question. It is possible that unit level causal effects change over time, so using a different data set with more recent data may simply tell us that the causal effect has been decreasing over time. If this is the case, your replication does not necessarily invalidate his analysis. On the other hand, if you have a strong reason to believe that the causal effects do not change over time, then this would perhaps call into question the validity of the earlier findings."
test_mc(correct = 2, feedback_msgs = c(msg1,msg2))
```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:1cc5a9ee77
## Causal Inference with Matching Methods 
Suppose you want to know the effect of enrollment in a Massive Open Online Course (MOOC) on income. You are worried that people with more traditional education-e.g., someone with a bachelor's degree-will be more likely to enroll in a MOOC and also more likely to make more income, compared to someone with less traditional education-e.g., someone with only a high school diploma. So you make the unconfoundedness assumption that, holding traditional education fixed, people randomly enroll in MOOCs. You then use matching to estimate the ATE of MOOC enrollment, and find that it is positive. What policy might you recommend from this, if any?

*** =instructions
- No recommendation-correlation does not equal causation and so we cannot conclude what the effect of MOOC enrollment is.
- MOOC enrollment should be promoted as a means of improving people's wages.
- The costs of enrolling in MOOCs should qualify for federal student loans, just as traditional educational approaches do.
- MOOCs should be banned so as to promote traditional educational methods.

*** =sct
```{r}
msg1 = "Assuming you believe the unconfoundedness assumption, the matching method does actually get a true causal effect. As usual, there is always room for debate on whether this assumption is credible, but this answer overstates that case."
msg2 = "Correct! Since the effect of MOOCs on income is positive, policies which promote MOOCs will tend to increase income on average."
msg3 = "Our evidence suggests that MOOCs may help individuals achieve better incomes, but we do not have enough informtion to justify whether MOOCs should qualify for student loans. For example, given that MOOCs are relatively inexpensive to enroll in, it may be unnecessary or unreasonably costly to create a system that provides federal student loans for MOOCs."
msg4 = "Since the study finds a positive treatment effect of MOOCs on income, this is an unintuitive conclusion to arrive at."
test_mc(correct = 2, feedback_msgs = c(msg1,msg2,msg3,msg4))
```



--- type:VideoExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:a731479ef2
## The Effect of Volunteer Military Service on Lifetime Earnings
*** =video_link
//player.vimeo.com/video/217555588


--- type:NormalExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:ef8cd14ae9
## Propensity Score Matching in R
Some sort of matching question

*** =instructions
- Do something

*** =pre_exercise_code
```{r}
install.packages("MatchIt")
library(MatchIt)

x<-rnorm(10,10)

```
*** =sample_code
```{r}
# Useful guide
```
*** =solution
```{r}
# Solution and ideal way to solve problem. For example, imagine we want the user to add 2+2 and store result in "x"
Solution1<-sum(x)
```
*** =sct
```{r}
test_object("Solution1")
success_msg("Good work!")
```



--- type:VideoExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:c6f282fde8
## The Credibility of the Unconfoundedness Assumption
*** =video_link
//player.vimeo.com/video/217555887
